<a id="home"></a>
## Assignment 1 - Gaussian Naive Bayes implementation
In this assignment you will need to implement different functions with a focus on Scaling<br/>

<b>Index of functions you need to implement (total possible points: 100)</b>:<br/>

| Section | Section-name | points | Section | Section-name | points | Section | Section-name | points | Section | Section-name | points |
| :- | :- | :- | :- | :- | :- | :- | :- | :- | :- | :- | :- | 
| i. | [Review Materials](#review_materials) | | ii. | [Gen. instructions](#gen_instructions) | | 1. | [Import packages](#Imports) | | 2. | [Student info](#student_info) | 5 | 
| 3. | [Load the dataset](#load_dataset) | 10 | 4.a. | [Calc. class priors](#calc_class_priors) | 10 | 4.b. | [Calc. mean for likelihood](#calc_mean) | 10 | 4.c. | [Calc. std for likelihood](#calc_std) |10 | 
| 4.d. | [fit (training)](#fit) | 10 | 5.a. | [Calc. PDF](#calc_gaussian_pdf_prob) | | 5.b. | [Calc. a-posterior](#calc_aposterior) | 15 | 5.c. | [predict](#predict) | 20 | 
| 6. | [Calc. error-rate](#calc_error_rate) | 10 |

<a id="review_materials"></a>
### (i) Materials to review for this assignment
<h4>From Moodle:</h4> 
<h5><u>Review the notebooks regarding the following python topics</u>:</h5>
<div class="alert alert-info">
&#x2714; Brief review: <b>Basic Python</b>.<br/>
&#x2714; Brief review: <b>The NumPy module</b><br/>
&#x2714; Thorough review: <b>The Pandas module</b> (both parts)<br/>
</div> 
<h5><u>Review the presentations regarding the following topics</u>:</h5>
<div class="alert alert-info">
&#x2714; <b>The Naive Bayes presentation</b><br/>
</div>

[Go to the beginning of the assignment](#home)
<a id="gen_instructions"></a>
### (ii) General assignment instructions
#### Working on the assignment, todo and not todo:
<b><u>Todo</u></b>:<br />
<div class="alert alert-success">
<ul>
<li>&#x2714; Use <b> Jupyter Notebook </b>.</li>
<li>&#x2714; <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students</li>
<li>&#x2714; Test and save your assignment - <b>submit the last tested and saved version</b></li>
<li>&#x2714; Work on the <b>original notebook</b> of the assignment</li>
<li>&#x2714; Run the code cells by their original order</li>
<li>&#x2714; submit <b>only</b> the <b>notebook (Assignment1_GaussianNaiveBayes.ipynb)</b> file (no name changing)</li></ul>
</div>
<b><u>Not todo</u></b>:<br />
<div class="alert alert-danger">
<ul>
<li>&#x2718; <b>do NOT use Google Collaboratory</b> </li>
<li>&#x2718; do <b>NOT</b> submit an <b>empty assignment</b></li>
<li>&#x2718; do <b>NOT submit extra files</b>, <u>unless you're asked to do so</u></li>
<li>&#x2718; <b>Do NOT submit a .py/.txt/.rar/.zip (or any non (.ipynb) file)</b> <u>version for the notebook</u> of the assignment</li>
<li>&#x2718; <b>do NOT change</b> the notebook file name</li>
</ul>
</div>

<b>Your implementations are performed in the 'Answer code cell's as following</b>:<br/>
<div class="alert alert-success">
&#x2714; Your code starts one line after the '<code># YOUR CODE HERE</code>' remark<br/>
&#x2714; Implemented code<br/>
&#x2714; returned expected value<br/><br/>
<b><u>Your code should look like this</u><b>:</div>
<code># 5. function_name (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def method_to_implement():
    # YOUR CODE HERE
    your implemented code ...
</code>

[Go to the beginning of the assignment](#home)
<a id="Imports"></a>
### 1. import modules (packages)
This step is necessary in order to use external modules (packages). <br/>
<div class="alert alert-success">
&#x2714; Any module or package imported in the following code cell are permitted
</div>
<div class="alert alert-danger">
&#x2718; Any other modules or packages are ILLEGAL
</div>

##### 1. IMPORT MODULES - Run the following code cell:

```python
# 1. IMPORT MODULES 
# ====>>> YOUR ACTIONS: RUN ONLY

import sys                       # for testing use only
import os                        # for testing use only
from datetime import datetime    # for testing use only
import random                    # for testing use only
import hashlib                   # for testing use only
import pandas as pd
import numpy as np
import math 
import statistics

# --------------------------------------
# show several outputs in one cell. 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"
# --------------------------------------
```

[Go to the beginning of the assignment](#home)
<a id="student_info"></a>
### 2. Student information
##### methods to implement:
* my_name
* my_id

#### 2.a. Function: my_name
<a id="myName"></a>
Implement the function to return your full name (as a string), for example:<br/>
<code>    return 'John Smith'</code>

##### 2.a. my_name - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: solution2a-my_name, points: 0, grade: false, solution: true -->
```python
# 2.a. my_name (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def my_name():
    # YOUR CODE HERE
    return "Netanel Peduim_Yakir Komemi"
```

##### 2.a.  my_name - OPTIONAL SELF TEST -  Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-15667dc5f925d99b, points: 0, grade: false, solution: true -->
```python
# 2.a. my_name (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

student_name = my_name()
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

#### 2.b. Function: my_id
<a id="myName"></a>
Implement the function to return your ID number (as an integer number).<br/>
Note: the **id number should be equal to the one in Moodle**, for example:<br/>
<code>    return 132468579</code>

##### 2.b. my_id - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: solution2b-my_id, points: 0, grade: false, solution: true -->
```python
# 2.b. my_id (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN 

def my_id():
    # YOUR CODE HERE
    return 203797816
```

##### 2.b.  my_id - OPTIONAL SELF TEST -  Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-1d5f9d602b1248de, points: 0, grade: false, solution: true -->
```python
# 2.b. my_id (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

student_id = my_id()
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

##### 2. my_name, my_id - GRADED TESTS - Run the following code cell:

<!-- grade_id: test2-1_student_info, points: 5, grade: true, solution: false -->
```python
# 2. my_name, my_id (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("2. 'my_name' and 'my_id' - Test 1 (name: test2-1_student_info, points: 5)")

# dataframe for output:
dt1 = datetime.now()
try:
    student_name, student_id = my_name(), my_id()
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert type(student_name) is str or type(student_id) is int, "name is not a string or id is not an integer"

s_datetime =  [dt1.strftime('%Y-%m-%d %H:%M:%S'), datetime.now().strftime('%Y-%m-%d %H:%M:%S')]
df_output = pd.DataFrame({'value': [student_name, student_id], 'type': [type(student_name), type(student_id)], 'date_time': s_datetime},index=['student_name','student_id'])

print ("\nGood Job!\nYou've passed the test for the implementation of 'my_name' and 'my_id'  :-)")

print ('\nOutput dataframe:')
print ('-----------------')
df_output
```

Output:
```
2. 'my_name' and 'my_id' - Test 1 (name: test2-1_student_info, points: 5)

Good Job!
You've passed the test for the implementation of 'my_name' and 'my_id'  :-)

Output dataframe:
-----------------
                                    value           type            date_time
student_name  Netanel Peduim_Yakir Komemi  <class 'str'>  2024-08-28 22:04:31
student_id                      203797816  <class 'int'>  2024-08-28 22:04:31
```

### 3. Load dataset

[Go to the beginning of the assignment](#home)
<a id="load_dataset"></a>
#### 3. Load dataset
<u>method name</u>: <b>load_dataset</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 

------------

<pre>
input parameters:
- file_name - a path containing the csv file name. Use the 'file_name' as is
- category_col_name - the category column name
------------
The function should return a dataframe containing the feature vectors (X) and 
    a series containing the categories (y). </pre>
    
------------

The return statement should look like this:<br/>
<code>return X, y</code>

##### 3. load_dataset - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-c95feb6aba5b1b7e, points: 0, grade: false, solution: true -->
```python
# 3. load_dataset (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 3. load_dataset - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol3-load_dataset, points: 0, grade: false, solution: true -->
```python
# 3. load_dataset (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def load_dataset(file_name, category_col_name):
    # YOUR CODE HERE
    df = pd.read_csv(file_name)
    return df , df.get([category_col_name])
    
```

##### 3. load_dataset - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-6878b5d0024a8368, points: 0, grade: false, solution: true -->
```python
# 3. load_dataset (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X,y = load_dataset(file_name, category_col_name)
X.head(3)
# --- add additional code to check your code if needed:
# YOUR CODE HERE
print(X.shape)
print(y.shape)
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)  \
0                6.0               2.2                5.0               1.5   
1                6.4               2.8                5.6               2.1   
2                7.7               3.0                6.1               2.3   

   iris_type  
0          2  
1          2  
2          2  
(120, 5)
(120, 1)
```

##### 3. load_dataset - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test3_1-load_dataset, points: 1, grade: true, solution: false -->
```python
# 3. load_dataset (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("3.'load_dataset' - Test 1 (name: test3_1-load_dataset, points: 1) - Sanity")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'load_dataset' function implementation :-)")   
```

Output:
```
3.'load_dataset' - Test 1 (name: test3_1-load_dataset, points: 1) - Sanity
Good Job!
You've passed the 1st test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test3_2-load_dataset, points: 4, grade: true, solution: false -->
```python
# 3. load_dataset (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("3.'load_dataset' - Test 2 (name: test3_2-load_dataset, points: 4)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert type(X) is pd.DataFrame or type(y) is pd.Series, "Wrong type for feature vectors dataframe or for category series"
np.testing.assert_array_equal(X.index, y.index, 'X and y should have the same index')
    
print ("Good Job!\nYou've passed the 2nd test for the 'load_dataset' function implementation :-)")        
```

Output:
```
3.'load_dataset' - Test 2 (name: test3_2-load_dataset, points: 4)
Good Job!
You've passed the 2nd test for the 'load_dataset' function implementation :-)
```

<!-- grade_id: test3_3-load_dataset, points: 5, grade: true, solution: false -->
```python
# 3. load_dataset (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("3. 'load_dataset' - Test 3 (name: test3_3-load_dataset, points: 5)")

file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
assert X.shape == (30,4) and X.shape[0] == y.shape[0], 'Wrong shape for feature vector dataframe or category series'

print ("Good Job!\nYou've passed the 3rd test for the 'load_dataset' function implementation :-)")        
```

Output:
```
3. 'load_dataset' - Test 3 (name: test3_3-load_dataset, points: 5)
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp\ipykernel_21200\1415929043.py in <module>
     14     raise
     15 
---> 16 assert X.shape == (30,4) and X.shape[0] == y.shape[0], 'Wrong shape for feature vector dataframe or category series'
     17 
     18 print ("Good Job!\nYou've passed the 3rd test for the 'load_dataset' function implementation :-)")

AssertionError: Wrong shape for feature vector dataframe or category series
```

### Part 4 - fit (training a Gaussian Naive Bayes model) 
  * Calculate class priors
  * Calculate the 'mean' element out of likelihood probability
  * Calculate the 'std' element out of likelihood probability
  * fit

[Go to the beginning of the assignment](#home)
<a id="calc_class_priors"></a>
#### 4.a. Calculate class priors 
<u>method name</u>: <b>calc_class_priors</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 
* <b>load_dataset</b></br> 

------------

<pre>
Change the following function to return a python list 
containing the 3 classes priors (one for each class).
------------
input parameter:
- y_train - a series containing the train-set's categories
- lst_class_vals - A simple python list, containing the possible category values. 
               remark: for the iris case it includes: [0,1,2]
------------
return value:
-  a python list, containing the class priors
</pre>
    
------------

The return statement should look like this:<br/>
<code>return lst_priors</code>

##### 4.a. calc_class_priors - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-934955b2d7946121, points: 0, grade: false, solution: true -->
```python
# 4.a. calc_class_priors (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 4.a. calc_class_priors - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol4a-calc_class_priors, points: 0, grade: false, solution: true -->
```python
# 4.a. calc_class_priors (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def calc_class_priors(y_train, lst_class_vals):
    # YOUR CODE HERE
    lst_priors = []
    all_values = len(y_train)
    for class_val in lst_class_vals:
        count = (y_train.values == class_val).sum()
        prior = count / all_values
        lst_priors.append(prior)
    return lst_priors
    
    
```

##### 4.a. calc_class_priors - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-f32089e054525334, points: 0, grade: false, solution: true -->
```python
# 4.a. calc_class_priors (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
lst_class_values = [0,1,2]
student_name, student_id = my_name(), my_id()
X, y = load_dataset(file_name, category_col_name)
lst_priors = calc_class_priors(y, lst_class_values)
# --- add additional code to check your code if needed:
# YOUR CODE HERE
print(np.round(lst_priors[0],2))
print(lst_priors)
```

Output:
```
0.34
[0.3416666666666667, 0.325, 0.3333333333333333]
```

##### 4.a. calc_class_priors - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test4a_1-calc_class_priors, points: 1, grade: true, solution: false -->
```python
# 4.a. calc_class_priors (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.a.'calc_class_priors' - Test 1 (name: test4a_1-calc_class_priors, points: 1) - Sanity")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
lst_class_values = [0,1,2]
try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    lst_priors = calc_class_priors(y, lst_class_values)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'calc_class_priors' function implementation :-)")   
```

Output:
```
4.a.'calc_class_priors' - Test 1 (name: test4a_1-calc_class_priors, points: 1) - Sanity
Good Job!
You've passed the 1st test for the 'calc_class_priors' function implementation :-)
```

<!-- grade_id: test4a_2-calc_class_priors, points: 4, grade: true, solution: false -->
```python
# 4.a. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.a.'calc_class_priors' - Test 2 (name: test4a_2-calc_class_priors, points: 4)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
lst_class_values = [0,1,2]

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    lst_priors = calc_class_priors(y, lst_class_values)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert len(lst_priors)==3,'there should be 3 class priors'
    
print ("Good Job!\nYou've passed the 2nd test for the 'calc_class_priors' function implementation :-)")        
```

Output:
```
4.a.'calc_class_priors' - Test 2 (name: test4a_2-calc_class_priors, points: 4)
Good Job!
You've passed the 2nd test for the 'calc_class_priors' function implementation :-)
```

<!-- grade_id: test4a_3-calc_class_priors, points: 5, grade: true, solution: false -->
```python
# 4.a. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.a.'calc_class_priors' - Test 3 (name: test4a_3-calc_class_priors, points: 5)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
lst_class_values = [0,1,2]

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    lst_priors = calc_class_priors(y, lst_class_values)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert 0.34==np.round(lst_priors[0],2),'wrong prior for class-0'
    
print ("Good Job!\nYou've passed the 3rd test for the 'calc_class_priors' function implementation :-)")        
```

Output:
```
4.a.'calc_class_priors' - Test 3 (name: test4a_3-calc_class_priors, points: 5)
Good Job!
You've passed the 3rd test for the 'calc_class_priors' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="calc_mean"></a>
#### 4.b. Calculate the 'mean' element out of likelihood probability
<u>method name</u>: <b>calc_mean_for_likelihood</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 
* <b>load_dataset</b></br> 

------------

<pre>
Change the following function to return mean of the values for each feature
this should be done separately for each class value (in the training set)
    * create a df_mean_likelihood, including the features as columns and classes as indexes
    * the cells contain the mean values for each class and feature 
------------
input parameter:
- training_set - the train-set dataframe including also the target value
- category_col_name - the category column name
------------
return value:
-  a dataframe, consisting of the feature names as columns and classes as indexes
   the cells consists of the mean values for each class and feature 
</pre>
    
------------

The return statement should look like this:<br/>
<code>return df_mean_likelihood</code>

##### 4.b. calc_mean_for_likelihood - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

```python
X.head(3)
y.head(3)
trainset=X.copy(deep=True)
trainset['iris_type']=y
trainset.head(3)
```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)  \
0                6.0               2.2                5.0               1.5   
1                6.4               2.8                5.6               2.1   
2                7.7               3.0                6.1               2.3   

   iris_type  
0          2  
1          2  
2          2  
   iris_type
0          2
1          2
2          2
   sepal length (cm)  sepal width (cm)  petal length (cm)  petal width (cm)  \
0                6.0               2.2                5.0               1.5   
1                6.4               2.8                5.6               2.1   
2                7.7               3.0                6.1               2.3   

   iris_type  
0          2  
1          2  
2          2
```

<!-- grade_id: cell-1aef425a6e8a4558, points: 0, grade: false, solution: true -->
```python
# 4.b. calc_mean_for_likelihood (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 4.b. calc_mean_for_likelihood - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol4b-calc_mean_for_likelihood, points: 0, grade: false, solution: true -->
```python
# 4.b. calc_mean_for_likelihood (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def calc_mean_for_likelihood(training_set, category_col_name):
    # YOUR CODE HERE
    grouped_means = training_set.groupby(category_col_name).mean()
    grouped_means['iris_type'] = [0,1,2] 
    print(grouped_means)
    return(grouped_means)
    
```

##### 4.b. calc_mean_for_likelihood - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-22f39df08f8af753, points: 0, grade: false, solution: true -->
```python
# 4.b. calc_mean_for_likelihood (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X, y = load_dataset(file_name, category_col_name)
trainset=X.copy(deep=True)
trainset[category_col_name]=y
df_mean_likelihood = calc_mean_for_likelihood(trainset, category_col_name)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type                               
0                  0.241463          0  
1                  1.328205          1  
2                  2.012500          2
```

##### 4.b. calc_mean_for_likelihood - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test4b_1-calc_mean_for_likelihood, points: 1, grade: true, solution: false -->
```python
# 4.b. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.b.'calc_mean_for_likelihood' - Test 1 (name: test4b_1-calc_mean_for_likelihood, points: 1) - Sanity")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    df_mean_likelihood = calc_mean_for_likelihood(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'calc_mean_for_likelihood' function implementation :-)")   
```

Output:
```
4.b.'calc_mean_for_likelihood' - Test 1 (name: test4b_1-calc_mean_for_likelihood, points: 1) - Sanity
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type                               
0                  0.241463          0  
1                  1.328205          1  
2                  2.012500          2  
Good Job!
You've passed the 1st test for the 'calc_mean_for_likelihood' function implementation :-)
```

<!-- grade_id: test4b_2-calc_mean_for_likelihood, points: 4, grade: true, solution: false -->
```python
# 4.b. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.b.'calc_mean_for_likelihood' - Test 2 (name: test4b_2-calc_mean_for_likelihood, points: 4)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    df_mean_likelihood = calc_mean_for_likelihood(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

np.testing.assert_array_equal(df_mean_likelihood.index, [0,1,2], 'bad index for df_mean_likelihood, should be equal to class names')
np.testing.assert_array_equal(df_mean_likelihood.columns, X.columns, "bad colmn names for df_mean_likelihood, should be equal to feature names (X's columns)")
    
print ("Good Job!\nYou've passed the 2nd test for the 'calc_mean_for_likelihood' function implementation :-)")        
```

Output:
```
4.b.'calc_mean_for_likelihood' - Test 2 (name: test4b_2-calc_mean_for_likelihood, points: 4)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type                               
0                  0.241463          0  
1                  1.328205          1  
2                  2.012500          2  
Good Job!
You've passed the 2nd test for the 'calc_mean_for_likelihood' function implementation :-)
```

<!-- grade_id: test4b_3-calc_mean_for_likelihood, points: 5, grade: true, solution: false -->
```python
# 4.b. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.b.'calc_mean_for_likelihood' - Test 3 (name: test4b_3-calc_mean_for_likelihood, points: 5)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    df_mean_likelihood = calc_mean_for_likelihood(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert 6.5==np.round(df_mean_likelihood.iloc[2,0],1),"Wrong mean value for 'sepal length' for class 2"
print ("Good Job!\nYou've passed the 3rd test for the 'calc_mean_for_likelihood' function implementation :-)")        
```

Output:
```
4.b.'calc_mean_for_likelihood' - Test 3 (name: test4b_3-calc_mean_for_likelihood, points: 5)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type                               
0                  0.241463          0  
1                  1.328205          1  
2                  2.012500          2  
Good Job!
You've passed the 3rd test for the 'calc_mean_for_likelihood' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="calc_std"></a>
#### 4.c. Calculate the 'std' element out of likelihood probability
<u>method name</u>: <b>calc_std_for_likelihood</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 
* <b>load_dataset</b></br> 

------------

<pre>
Change the following function to return standard deviation (std) of the values for each feature
this should be done separately for each class value (in the training set)
    * create a df_std_likelihood, including the features as columns and classes as indexes
    * the cells contain the standard deviation (std) values for each class and feature 
------------
input parameter:
- training_set - the train-set dataframe including also the target value
- category_col_name - the category column name
------------
return value:
-  a dataframe, consisting of the feature names as columns and classes as indexes
   the cells contain the standard deviation (std) values for each class and feature 
</pre>
    
------------

The return statement should look like this:<br/>
<code>return df_std_likelihood</code>

##### 4.c. calc_std_for_likelihood - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-75cf95e220b98450, points: 0, grade: false, solution: true -->
```python
# 4.c. calc_std_for_likelihood (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 4.c. calc_std_for_likelihood - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol4c-calc_std_for_likelihood, points: 0, grade: false, solution: true -->
```python
# 4.c. calc_std_for_likelihood (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def calc_std_for_likelihood(training_set, category_col_name):
    # YOUR CODE HERE
    grouped_stds = training_set.groupby(category_col_name).std()
    grouped_stds['iris_type'] = [0,1,2] 
    print(grouped_stds)
    return(grouped_stds)
```

##### 4.c. calc_std_for_likelihood - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-21811508fed055de, points: 0, grade: false, solution: true -->
```python
# 4.c. calc_std_for_likelihood (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X, y = load_dataset(file_name, category_col_name)
trainset=X.copy(deep=True)
trainset[category_col_name]=y
df_std_likelihood = calc_std_for_likelihood(trainset,category_col_name)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2
```

##### 4.c. calc_std_for_likelihood - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test4c_1-calc_std_for_likelihood, points: 1, grade: true, solution: false -->
```python
# 4.c. calc_std_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.c.'calc_std_for_likelihood' - Test 1 (name: test4c_1-calc_std_for_likelihood, points: 1) - Sanity")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    df_std_likelihood = calc_std_for_likelihood(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'calc_std_for_likelihood' function implementation :-)")   
```

Output:
```
4.c.'calc_std_for_likelihood' - Test 1 (name: test4c_1-calc_std_for_likelihood, points: 1) - Sanity
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
Good Job!
You've passed the 1st test for the 'calc_std_for_likelihood' function implementation :-)
```

<!-- grade_id: test4c_2-calc_std_for_likelihood, points: 4, grade: true, solution: false -->
```python
# 4.c. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.c.'calc_std_for_likelihood' - Test 2 (name: test4c_2-calc_std_for_likelihood, points: 4)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    df_std_likelihood = calc_std_for_likelihood(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

np.testing.assert_array_equal(df_std_likelihood.index, [0,1,2], 'bad index for df_std_likelihood, should be equal to class names')
np.testing.assert_array_equal(df_std_likelihood.columns, X.columns, "bad colmn names for df_std_likelihood, should be equal to feature names (X's columns)")
    
print ("Good Job!\nYou've passed the 2nd test for the 'calc_std_for_likelihood' function implementation :-)")        
```

Output:
```
4.c.'calc_std_for_likelihood' - Test 2 (name: test4c_2-calc_std_for_likelihood, points: 4)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
Good Job!
You've passed the 2nd test for the 'calc_std_for_likelihood' function implementation :-)
```

<!-- grade_id: test4c_3-calc_std_for_likelihood, points: 5, grade: true, solution: false -->
```python
# 4.c. calc_mean_for_likelihood (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.c.'calc_std_for_likelihood' - Test 3 (name: test4c_3-calc_std_for_likelihood, points: 5)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    df_std_likelihood = calc_std_for_likelihood(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert 0.501==np.round(df_std_likelihood.iloc[1,2],3),"Wrong mean value for 'petal length' for class 1"
    
print ("Good Job!\nYou've passed the 3rd test for the 'calc_std_for_likelihood' function implementation :-)")        
```

Output:
```
4.c.'calc_std_for_likelihood' - Test 3 (name: test4c_3-calc_std_for_likelihood, points: 5)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
Good Job!
You've passed the 3rd test for the 'calc_std_for_likelihood' function implementation :-)
```

[Go to the beginning of the assignment](#home)
<a id="fit"></a>
#### 4.d. Gaussian Naive Bayes main training flow
<u>method name</u>: <b>fit</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 
* <b>load_dataset</b></br> 
* <b>calc_class_priors</b></br> 
* <b>calc_mean_for_likelihood</b></br> 
* <b>calc_std_for_likelihood</b></br> 

------------

<pre>
Change the following function to return the 3 products of the Gaussian Naive Bayes training.
The method needs to calculate and return the products from the 3 previous methods you've 
implemented in the previous cells, including calc_category_priors, calc_mean_for_likelihood, calc_std_for_likelihood.
Please check the expected output for these methods.
------------
input parameter:
- training_set - the train-set dataframe including also the target value
- category_col_name - the category column name
------------
return values (comma seperated):
- lst_class_prior    - output from calc_category_priors
- df_mean_likelihood - output from calc_mean_likelihood
- df_std_likelihood  - output from calc_std_likelihood
</pre>
    
------------

The return statement should look like this:<br/>
<code>return lst_class_prior,df_mean_likelihood,df_std_likelihood</code>

##### 4.d. fit - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-8ba5cac83a654b63, points: 0, grade: false, solution: true -->
```python
# 4.d. fit (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 4.d. fit - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol4d-fit, points: 0, grade: false, solution: true -->
```python
# 4.d. fit (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def fit(training_set,category_col_name):
    # YOUR CODE HERE
    df_std_likelihood = calc_std_for_likelihood(training_set,category_col_name)
    df_mean_likelihood = calc_mean_for_likelihood(training_set,category_col_name)
    lst_class_prior = calc_class_priors(training_set,[0,1,2])
    
    return lst_class_prior,df_mean_likelihood,df_std_likelihood
```

##### 4.d. fit - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-a916eb1dd94353b6, points: 0, grade: false, solution: true -->
```python
# 4.d. fit (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X, y = load_dataset(file_name, category_col_name)
trainset=X.copy(deep=True)
trainset[category_col_name]=y
lst_class_prior,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type                               
0                  0.241463     
```

##### 4.d. fit - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test4d_1-fit, points: 1, grade: true, solution: false -->
```python
# 4.d. fit (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.d.'fit' - Test 1 (name: test4d_1-fit, points: 1) - Sanity")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    lst_class_prior,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'fit' function implementation :-)")   
```

Output:
```
4.d.'fit' - Test 1 (name: test4d_1-fit, points: 1) - Sanity
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
ir
```

<!-- grade_id: test4c_2-fit, points: 4, grade: true, solution: false -->
```python
# 4.d. fit (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.d.'fit' - Test 2 (name: test4c_2-fit, points: 4)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    lst_class_prior,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert len(lst_priors)==3,'there should be 3 class priors'

np.testing.assert_array_equal(df_mean_likelihood.index, [0,1,2], 'bad index for df_mean_likelihood, should be equal to class names')
np.testing.assert_array_equal(df_mean_likelihood.columns, X.columns, "bad colmn names for df_mean_likelihood, should be equal to feature names (X's columns)")

np.testing.assert_array_equal(df_std_likelihood.index, [0,1,2], 'bad index for df_std_likelihood, should be equal to class names')
np.testing.assert_array_equal(df_std_likelihood.columns, X.columns, "bad colmn names for df_std_likelihood, should be equal to feature names (X's columns)")

    
print ("Good Job!\nYou've passed the 2nd test for the 'fit' function implementation :-)")        
```

Output:
```
4.d.'fit' - Test 2 (name: test4c_2-fit, points: 4)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type  
```

<!-- grade_id: test4c_3-fit, points: 5, grade: true, solution: false -->
```python
# 4.d. fit (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("4.d.'fit' - Test 3 (name: test4c_3-fit, points: 5)")

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'

try:
    student_name, student_id = my_name(), my_id()
    X, y = load_dataset(file_name, category_col_name)
    trainset=X.copy(deep=True)
    trainset[category_col_name]=y
    lst_class_prior,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert 0.34==np.round(lst_priors[0],2),'wrong prior for class-0'
assert 6.5==np.round(df_mean_likelihood.iloc[2,0],1),"Wrong mean value for 'sepal length' for class 2"
assert 0.501==np.round(df_std_likelihood.iloc[1,2],3),"Wrong mean value for 'petal length' for class 1"
    
print ("Good Job!\nYou've passed the 3rd test for the 'fit' function implementation :-)")        
```

Output:
```
4.d.'fit' - Test 3 (name: test4c_3-fit, points: 5)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type  
```

### Part 5 - predict 
  * Calculate Guassian likelihood probability 
  * Calculate a posteriori probabilities 
  * prdict

[Go to the beginning of the assignment](#home)
<a id="calc_gaussian_pdf_prob"></a>
#### 5.a. Calculate Gaussian likelihood PDF (instead of dependent prob. for continues variables)

<u>method name</u> <b>calc_gaussian_pdf_prob</b>

------------

<pre>
The following function uses the mean & std likelihood values to return the Gaussian PDF 
    (which replaces the probablilty function for continues variables).
The PDF calculated replaces the dependent probability such as pr(x_j|c), where x_j is a value for the
    feature x_j and c is a specific values of a class.
------------
input parameter:
- x_feature_val - a scalar value for a specific feature
- feature_mean  - the mean of the feature (in the training data) for a specific class
- feature_std   - the std of the feature (in the training data) for a specific class
------------
return value:
- pdf_val - the pdf calculated value
</pre>

##### 5.a. Calculate Gaussian likelihood PDF - Run the following code cell:

```python
# 5.a. calc_gaussian_pdf_prob (Calculate "dependent" Gaussian likelihood PDF) 
# ====>>> YOUR ACTIONS: RUN ONLY

def calc_gaussian_pdf_prob(x_feature_val, feature_mean, feature_std):
    exponent = np.exp(-((x_feature_val-feature_mean)**2 / (2 * feature_std**2 )))
    return (1 / ((2 * np.pi)**(1/2) * feature_std)) * exponent
```

[Go to the beginning of the assignment](#home)
<a id="calc_aposterior"></a>
#### 5.b. Calculate a-posterior probabilities
<u>method name</u>: <b>calc_aposterior_probs</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 
* <b>load_dataset</b></br> 

For tests:
* <b>calc_class_priors</b></br> 
* <b>calc_mean_for_likelihood</b></br> 
* <b>calc_std_for_likelihood</b></br>  
* <b>fit</b></br> 

------------

<pre>
The following function should create a probability matrix to store the results.
To do that create a dataframe with the categories as columns and test indexes ans indexes.
The value of the cells should be a multiplication between the class prior and the Gaussian probabilities of all the features (using the  'calc_gaussian_pdf_prob' method, for each pr(x[j]|c) probability) for a specific test feature vector (i.e., Gaussian probabilities for feature values of every row).
The outcome is the a-posterior probabilities for each feature vector (one for each class). 
Please check the expected output for these methods.
------------
input parameters:
- X_test - the test-set feature-vectors dataframe
- lst_class_priors - the trained class priors
- df_trained_mean - the dataframe of the trained means for each feature per each class
- df_trained_std - the dataframe of the trained std values, for each feature per each class
- lst_class_vals - A simple python list, containing the possible category values. 
               remark: for the iris case it includes: [0,1,2]
------------
return value:
- df_aposterior_probs -  a dataframe containing a-posterior probabilities per class, per test feature vector 
</pre>
    
------------

The return statement should look like this:<br/>
<code>return df_aposterior_probs</code>

##### 5.b. calc_aposterior_probs - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-91a0076514015ec1, points: 0, grade: false, solution: true -->
```python
# 5.b. calc_aposterior_probs (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 5.b. calc_aposterior_probs - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol5b-calc_aposterior_probs, points: 0, grade: false, solution: true -->
```python
# 5.b. calc_aposterior_probs (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def calc_aposterior_probs(X_test, lst_class_priors, df_trained_mean, df_trained_std, lst_class_vals):
    num_classes = len(lst_class_vals)
    df_prob_per_test_inst_per_class = pd.DataFrame(np.zeros((X_test.shape[0], num_classes)), columns=lst_class_vals, index=X_test.index)
    # YOUR CODE HERE
    
    
    
    
```

##### 5.b. calc_aposterior_probs - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-02f5ffa8a9091a1e, points: 0, grade: false, solution: true -->
```python
# 5.b. calc_aposterior_probs (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X_train, y_train = load_dataset(train_file_name, category_col_name)
unique_categories = np.sort(np.unique(y_train.values))
trainset=X_train.copy(deep=True)
trainset[category_col_name]=y_train
X_test, y_test = load_dataset(test_file_name, category_col_name)
lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
df_aposterior_probs = calc_aposterior_probs(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
# --- add additional code to check your code if needed:
# YOUR CODE HERE


```

Output:
```
sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal width (cm)  iris_type  
iris_type                               
0                  0.241463     
```

##### 5.b. calc_aposterior_probs - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test5b_1-calc_aposterior_probs, points: 1, grade: true, solution: false -->
```python
# 5.b. calc_aposterior_probs (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.b.'calc_aposterior_probs' - Test 1 (name: test5b_1-calc_aposterior_probs, points: 1) - Sanity")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    df_aposterior_probs = calc_aposterior_probs(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'calc_aposterior_probs' function implementation :-)")   
```

Output:
```
5.b.'calc_aposterior_probs' - Test 1 (name: test5b_1-calc_aposterior_probs, points: 1) - Sanity
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

       
```

<!-- grade_id: test5b_2-calc_aposterior_probs, points: 6, grade: true, solution: false -->
```python
# 5.b. calc_aposterior_probs (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.b.'calc_aposterior_probs' - Test 2 (name: test5b_2-calc_aposterior_probs, points: 6)")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    df_aposterior_probs = calc_aposterior_probs(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

np.testing.assert_array_equal (df_aposterior_probs.columns, unique_categories, "Wrong columns for 'df_aposterior_probs', should be equals to the three classes: 1,2,3")
np.testing.assert_array_equal(df_aposterior_probs.index, X_test.index, 'df_aposterior_probs and X_test should have the same index')
    
print ("Good Job!\nYou've passed the 2nd test for the 'calc_aposterior_probs' function implementation :-)")        
```

Output:
```
5.b.'calc_aposterior_probs' - Test 2 (name: test5b_2-calc_aposterior_probs, points: 6)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal
```

<!-- grade_id: test5b_3-calc_aposterior_probs, points: 8, grade: true, solution: false -->
```python
# 5.b. calc_aposterior_probs (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.b.'calc_aposterior_probs' - Test 3 (name: test5b_3-calc_aposterior_probs, points: 8)")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    df_aposterior_probs = calc_aposterior_probs(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert np.round(df_aposterior_probs.iloc[14,0],1)==1.6,"Wrong value for a-posterior probability for for 15th example (index 14) and class 0"
    
print ("Good Job!\nYou've passed the 3rd test for the 'calc_aposterior_probs' function implementation :-)")        
```

Output:
```
5.b.'calc_aposterior_probs' - Test 3 (name: test5b_3-calc_aposterior_probs, points: 8)
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   0.376100          0.395724           0.162938   
1                   0.499271          0.322712           0.501146   
2                   0.599551          0.314561           0.528932   

           petal width (cm)  iris_type  
iris_type                               
0                  0.104823          0  
1                  0.212688          1  
2                  0.270979          2  
           sepal length (cm)  sepal width (cm)  petal length (cm)  \
iris_type                                                           
0                   4.982927          3.412195           1.446341   
1                   5.961538          2.782051           4.279487   
2                   6.505000          2.945000           5.465000   

           petal
```

[Go to the beginning of the assignment](#home)
<a id="predict"></a>
#### 5.c. predict - classify new examples
<u>method name</u>: <b>predict</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 
* <b>load_dataset</b></br> 
* <b>calc_aposterior_probs</b></br> 

For tests:
* <b>calc_class_priors</b></br> 
* <b>calc_mean_for_likelihood</b></br> 
* <b>calc_std_for_likelihood</b></br>  
* <b>fit</b></br> 

------------

<pre>
Change the following function to return the predicted categories of the test set examples
one category for each corresponding test instance
this should be done separately for each class value (in the training set), using the Gaussian Naive Bayes.
Note: For each test feature vector, use the a-posterior probability for each class. 
      The predicted class (in the example level), should be the one with the maximum probability.
------------
input parameter:
- X_test - the test-set instances
- lst_class_priors - the trained priors of the classes
- df_trained_mean - the dataframe of the trained means for each feature per each class
- df_trained_std - the dataframe of the trained stds for each feature per each class
- lst_class_vals - a simple python list, containing the possible classes.
               remark: for the iris case it includes: [0,1,2]
------------
return value:
- y_predicted -  a series containing all predicted class values per test example.
                 Note: the index in the series should correspond to the X_test indexes
</pre>

------------

The return statement should look like this:<br/>
<code>return y_predicted</code>

##### 5.c. predict - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-0333d91f23c92573, points: 0, grade: false, solution: true -->
```python
# 5.c. predict (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 5.c. predict  - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol5c-predict, points: 0, grade: false, solution: true -->
```python
# 5.c. predict (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def predict(X_test, lst_class_priors, df_trained_mean, df_trained_std, lst_class_vals):
    # YOUR CODE HERE
    
```

##### 5.c. predict - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-8d65f5e2672c7f35, points: 0, grade: false, solution: true -->
```python
# 5.c. predict (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X_train, y_train = load_dataset(train_file_name, category_col_name)
unique_categories = np.sort(np.unique(y_train.values))
trainset=X_train.copy(deep=True)
trainset[category_col_name]=y_train
X_test, y_test = load_dataset(test_file_name, category_col_name)
lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
y_predict=predict(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

##### 5.c. predict - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test5c_1-predict, points: 1, grade: true, solution: false -->
```python
# 5.c. predict (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.c.'predict' - Test 1 (name: test5c_1-predict, points: 1) - Sanity")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    y_predict=predict(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'predict' function implementation :-)")   
```

<!-- grade_id: test5c_2-predict, points: 5, grade: true, solution: false -->
```python
# 5.b. calc_aposterior_probs (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.c.'predict' - Test 2 (name: test5c_2-predict, points: 5)")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    y_predicted=predict(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
assert y_predicted.shape[0]==y_test.shape[0],'predict should return a series with %d rows' %(y_test.shape[0])
np.testing.assert_array_equal(y_predicted.index, y_test.index, 'y_predicted and y_test should have the same index')

print ("Good Job!\nYou've passed the 2nd test for the 'predict' function implementation :-)")        
```

<!-- grade_id: test5c_3-predict, points: 7, grade: true, solution: false -->
```python
# 5.c. calc_aposterior_probs (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.c.'predict' - Test 3 (name: test5c_3-predict, points: 7)")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    y_predicted=predict(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert y_predicted.iloc[0]==0,'wrong predicted value for 1st value, expects: 0'
    
print ("Good Job!\nYou've passed the 3rd test for the 'predict' function implementation :-)")        
```

<!-- grade_id: test5c_4-predict, points: 7, grade: true, solution: false -->
```python
# 5.c. calc_aposterior_probs (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("5.c.'predict' - Test 4 (name: test5c_4-predict, points: 7)")

train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    y_predicted=predict(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert y_predicted.iloc[-1]==2,'wrong predicted value for last value, expects: 2'
    
print ("Good Job!\nYou've passed the 4th test for the 'predict' function implementation :-)")        
```

[Go to the beginning of the assignment](#home)
<a id="calc_error_rate"></a>
### 6. Evaluation (error rate)
<u>method name</u>: <b>calc_error_rate</b><br/>

------------

<u>dependency</u>: <br/>
* <b>my_name</b></br> 
* <b>my_id</b></br> 

------------

<pre>
Change the following function to return the error rate (incorrectly predicted ratio).
------------
input parameter:
- y_predicted -  a series of containing predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return value:
- errorRateVal - the error rate value
</pre>

------------

The return statement should look like this:<br/>
<code>return errorRateVal</code>

##### 6. calc_error_rate - OPTIONAL  ASSISTANCE  - Implement & Run the following cell (IF NEEDED):

<!-- grade_id: cell-8ba28f0677809d45, points: 0, grade: false, solution: true -->
```python
# 6. calc_error_rate (optional assitance for your solution)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

# ---- Add assistance code here IF NEEDED:
# YOUR CODE HERE

```

##### 6. calc_error_rate  - YOUR SOLUTION - Implement and Run the following code cell:

<!-- grade_id: sol6-calc_error_rate, points: 0, grade: false, solution: true -->
```python
# 6. calc_error_rate (your solution)
# ====>>> YOUR ACTIONS: IMPLEMENT & RUN

def calc_error_rate(y_predicted,y_test):
    # YOUR CODE HERE
    
```

##### 6. calc_error_rate - OPTIONAL SELF TEST - Implement and Run the following code cell (IF NEEDED):

<!-- grade_id: cell-caf45242507c06a6, points: 0, grade: false, solution: true -->
```python
# 6. calc_error_rate (optional self testing)
# ====>>> YOUR ACTIONS: (optional) IMPLEMENT & RUN 

file_name = 'iris_train_set.csv'
category_col_name = 'iris_type'
student_name, student_id = my_name(), my_id()
X, y = load_dataset(file_name, category_col_name)
# --- add additional code to check your code if needed:
# YOUR CODE HERE

```

##### 6. calc_error_rate - GRADED TESTS - Run the following (3) code cells:

<!-- grade_id: test6_1-calc_error_rate, points: 1, grade: true, solution: false -->
```python
# 6. calc_error_rate (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("6.'calc_error_rate' - Test 1 (name: test6_1-calc_error_rate, points: 1) - Sanity")

try:
    student_name, student_id = my_name(), my_id()
    y_test     =pd.Series([0,1,1,0,1,0,1,1,1,0])
    y_predicted=pd.Series([0,1,0,0,0,0,1,1,1,1])
    error_rate = calc_error_rate(y_predicted,y_test)    
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise
    
print ("Good Job!\nYou've passed the 1st test for the 'calc_error_rate' function implementation :-)")   
```

<!-- grade_id: test5b_2-calc_error_rate, points: 4, grade: true, solution: false -->
```python
# 6. calc_error_rate (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("6.'predict' - Test 2 (name: test5b_2-calc_error_rate, points: 4)")

try:
    student_name, student_id = my_name(), my_id()
    y_test     =pd.Series([0,1,1,0,1,0,1,1,1,0])
    y_predicted=pd.Series([0,1,0,0,0,0,1,1,1,1])
    error_rate = calc_error_rate(y_predicted,y_test)  
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert np.round(error_rate,2)==0.3,'Wrong error rate for test set'

print ("Good Job!\nYou've passed the 2nd test for the 'calc_error_rate' function implementation :-)")        
```

<!-- grade_id: test5b_3-calc_error_rate, points: 5, grade: true, solution: false -->
```python
# 6. calc_error_rate (graded test)
# ====>>> YOUR ACTIONS: RUN ONLY

print ("6.'predict' - Test 3 (name: test5b_3-calc_error_rate, points: 5)")


train_file_name = 'iris_train_set.csv'
test_file_name = 'iris_test_set.csv'
category_col_name = 'iris_type'
try:
    student_name, student_id = my_name(), my_id()
    X_train, y_train = load_dataset(train_file_name, category_col_name)
    unique_categories = np.sort(np.unique(y_train.values))
    trainset=X_train.copy(deep=True)
    trainset[category_col_name]=y_train
    X_test, y_test = load_dataset(test_file_name, category_col_name)
    lst_class_priors,df_mean_likelihood,df_std_likelihood = fit(trainset,category_col_name)
    y_predicted=predict(X_test, lst_class_priors, df_mean_likelihood, df_std_likelihood, unique_categories)
    error_rate = calc_error_rate(y_predicted,y_test)
except Exception as e:
    print ('You probably have a syntax or implementation error,  \nerror Message:',str(e), '\nTry fixing your code')
    raise

assert np.round(error_rate,2)==0.0,'Wrong error rate for test set'

print ("Good Job!\nYou've passed the 3rd test for the 'calc_error_rate' function implementation :-)")        
```

```python

```